# DS 653 -- Homework 1

**Due:** Friday, January 31 at 10pm on [Gradescope](https://www.gradescope.com/courses/959425).

_You must follow the Academic Code of Conduct and Collaboration Policy stated in the course syllabus at all times while working on this assignment._

This assignment contains 10 questions worth a total of 24 points. You will earn:
- 1 course learning credit for earning at least 10 points, or
- 2 course learning credits for earning at least 20 points on this homework.

To begin, please execute the code block below:

```python
# Initialize otter-grader
import otter
grader = otter.Notebook()

# Import libraries related to cryptography
import Cryptodome
from binascii import hexlify, unhexlify
from hashlib import sha256


# If this cell fails to execute, then use pip to install the following packages:
# otter-grader
# pycryptodomex
```

### Reading Questions (2 points each)

These multiple-choice questions are based on the week 1 and 2 reading assignments:

- David Wong’s [Hash functions and security](https://freecontent.manning.com/hash-functions-and-security/)
- [NBFMG, chapter 1](https://d28rh4a8wq0iu5.cloudfront.net/bitcointech/readings/princeton_bitcoin_book.pdf)

**Question 1.** For the hash function `SHA-256`, does the number 256 refer to the number of bits in the input or the output?
1. input
2. output

To answer this question, please uncomment out the correct response.

```python
def hw1_q1() -> int:
    # return 1
    # return 2
    return 2
```

```python
grader.check("q1")
```

Output:
```
File ~/ds653/.venv/lib/python3.12/site-packages/otter/check/notebook.py:256, in Notebook.check(self, question, global_env)
    243 """
    244 Runs tests for a specific question against a global environment. If no global environment
    245 is provided, the test is run against the calling frame's environment.
   (...)
    253     ``otter.test_files.abstract_test.TestFile``: the grade for the question
    254 """
    255 self._logger.info(f"Running check for question: {question}")
--> 256 test_path, test_name = resolve_test_info(
    257     self._tests_dir,
    258     self._notebook,
    259     self._tests_url_prefix,
    260     question,
    261 )
    263 self._logger.debug(f"Resolved test path: {test_path}")
    264 self._logger.debug(f"Resolved test name: {test_name}")

File ~/ds653/.venv/lib/python3.12/site-packages/otter/check/utils.py:351, in resolve_test_info(tests_dir, nb_path, tests_url_prefix, question)
    349 else:
    350     if nb_path is None:
--> 351         raise Va
```

**Question 2.** Which security property of cryptographic hash functions ensures that no one can reverse the hash function?

1. pre-image resistance
2. second pre-image resistance
3. collision resistance
4. random oracle

```python
def hw1_q2() -> int:
    # return 1
    # return 2
    # return 3
    # return 4
    return 1
```

```python
grader.check("q2")
```

Output:
```
File ~/ds653/.venv/lib/python3.12/site-packages/otter/check/notebook.py:256, in Notebook.check(self, question, global_env)
    243 """
    244 Runs tests for a specific question against a global environment. If no global environment
    245 is provided, the test is run against the calling frame's environment.
   (...)
    253     ``otter.test_files.abstract_test.TestFile``: the grade for the question
    254 """
    255 self._logger.info(f"Running check for question: {question}")
--> 256 test_path, test_name = resolve_test_info(
    257     self._tests_dir,
    258     self._notebook,
    259     self._tests_url_prefix,
    260     question,
    261 )
    263 self._logger.debug(f"Resolved test path: {test_path}")
    264 self._logger.debug(f"Resolved test name: {test_name}")

File ~/ds653/.venv/lib/python3.12/site-packages/otter/check/utils.py:351, in resolve_test_info(tests_dir, nb_path, tests_url_prefix, question)
    349 else:
    350     if nb_path is None:
--> 351         raise Va
```

**Question 3.** A cryptographic hash function can be used to construct a message commitment scheme. Which security property of the hash function prevents the committer from revealing a different message (after the commitment has been sent to the receiver)?

1. pre-image resistance
2. second pre-image resistance
3. collision resistance
4. random oracle

```python
def hw1_q3() -> int:
    # return 1
    # return 2
    # return 3
    # return 4
    return 3
```

```python
grader.check("q3")
```

Output:
```
File ~/ds653/.venv/lib/python3.12/site-packages/otter/check/notebook.py:256, in Notebook.check(self, question, global_env)
    243 """
    244 Runs tests for a specific question against a global environment. If no global environment
    245 is provided, the test is run against the calling frame's environment.
   (...)
    253     ``otter.test_files.abstract_test.TestFile``: the grade for the question
    254 """
    255 self._logger.info(f"Running check for question: {question}")
--> 256 test_path, test_name = resolve_test_info(
    257     self._tests_dir,
    258     self._notebook,
    259     self._tests_url_prefix,
    260     question,
    261 )
    263 self._logger.debug(f"Resolved test path: {test_path}")
    264 self._logger.debug(f"Resolved test name: {test_name}")

File ~/ds653/.venv/lib/python3.12/site-packages/otter/check/utils.py:351, in resolve_test_info(tests_dir, nb_path, tests_url_prefix, question)
    349 else:
    350     if nb_path is None:
--> 351         raise Va
```

**Question 4.** A digital signature scheme contains three algorithms: key generation, sign, and verify. Which is the only algorithm that requires the *secret key* in order to execute?

1. Key generation
2. Sign
3. Verify

```python
def hw1_q4() -> int:
    # return 1
    # return 2
    # return 3
    return 2
```

```python
grader.check("q4")
```

**Question 5.** A digital signature scheme contains three algorithms: key generation, sign, and verify. Which is the only algorithm that is guaranteed to be *deterministic* -- that is, if you run it multiple times with the same input, then you are guaranteed to receive the same response each time?

1. Key generation
2. Sign
3. Verify

```python
def hw1_q5() -> int:
    # return 1
    # return 2
    # return 3
    return 3
```

```python
grader.check("q5")
```

**Question 6.** Suppose that Alice and Bob use the ScroogeCoin system described in Section 1.5 of [NBFMG](https://d28rh4a8wq0iu5.cloudfront.net/bitcointech/readings/princeton_bitcoin_book.pdf). Suppose that Alice would like to transfer 1 of her ScroogeCoins to Bob. To accomplish this task, which person needs to create a digital signature?

1. Alice
2. Bob
3. Scrooge

```python
def hw1_q6() -> int:
    # return 1
    # return 2
    # return 3
    return 1
```

```python
grader.check("q6")
```

### Programming Questions (3 points each)

These questions require you to write code in `Python`. Many of this week's questions involve type conversions between int, string, and bytestring. The final question involves hash functions.

**Question 7 (Integer to Hex).** Write a function `int_byte_to_hex` that receives a Python integer in the range 0 to 255, and returns a Python string corresponding to the same number in base 16 / hexadecimal format.

As background: remember that a _byte_ of information contains 8 bits, so it can be used to represent a number between 0 and 255 (inclusive). In crypto, we often convert between different representations of a byte of information. One common representation that we will use often in this class is hexadecimal format, in which each symbol (0-9 or a-f) corresponds to a number between 0 and 15, and a byte can be represented with exactly two symbols.

```python
def int_byte_to_hex(the_input : int) -> str:
    """ Read a single, decimal byte from the user and return a string of its
    hexidecimal value. This string should use lowercase and should always be
    exactly two characters long. Make sure you pad the beginning with a 0 if
    necessary, and make sure the string does NOT start with '0x'.

    Example test cases:

        255 -> "ff"
        10 -> "0a"
        65 -> "41"
        161 -> "a1"
    """
    assert(type(the_input) == int)               # the input is an integer
    assert(the_input >= 0 and the_input <= 255)  # between 0 and 255, inclusive
    return hex(the_input)[2:].zfill(2)

    # todo: implement the rest of this function
```

As you will see throughout this course, crypto code can be incredibly precise and finicky. For this reason: most programming questions will be accommpanied by **test cases**, which are example inputs and outputs to a function that allow you to check whether your code is working properly.

In this question, we provide test cases in human-readable format in the comment above. We also provide unit tests for you to execute with `otter-grader`.

```python
grader.check("q7")
```

Output:
```
File ~/ds653/.venv/lib/python3.12/site-packages/otter/check/notebook.py:256, in Notebook.check(self, question, global_env)
    243 """
    244 Runs tests for a specific question against a global environment. If no global environment
    245 is provided, the test is run against the calling frame's environment.
   (...)
    253     ``otter.test_files.abstract_test.TestFile``: the grade for the question
    254 """
    255 self._logger.info(f"Running check for question: {question}")
--> 256 test_path, test_name = resolve_test_info(
    257     self._tests_dir,
    258     self._notebook,
    259     self._tests_url_prefix,
    260     question,
    261 )
    263 self._logger.debug(f"Resolved test path: {test_path}")
    264 self._logger.debug(f"Resolved test name: {test_name}")

File ~/ds653/.venv/lib/python3.12/site-packages/otter/check/utils.py:351, in resolve_test_info(tests_dir, nb_path, tests_url_prefix, question)
    349 else:
    350     if nb_path is None:
--> 351         raise Va
```

**Question 8 (Binary string to Hex).** Write a function `bytestring_to_hex_bytestring` that receives a Python bytestring as input, and returns the hexadecimal representation of the same bytestring.

A few hints:

- In Python, the prefix `b` changes the type of the Python variable from `string` to `bytes`. Since we will usually prefer to use `bytes` in this course, remember to always write the letter `b` before the quotation marks when creating bytestrings in this course! (Otherwise, you can find yourself stuck with an annoying error that is difficult to debug.)

```python
print(type("puzzle"), type(b"puzzle"))
```

Output:
```
<class 'str'> <class 'bytes'>
```

- Once again, we provide several example test cases for you in the function spec. (The otter-grader automatic grading system will test you on even more examples.)
- The `binascii` library includes a function that will do this for you. Take a look at its documentation here: https://docs.python.org/3/library/binascii.html

```python
def bytestring_to_hex_bytestring(the_input : bytes) -> bytes:
    """ Take in a string (byte object), and return the hex string of the bytes
    corresponding to it. While the hexlify() command will do this for you,
    we ask that you instead solve this question by combining the methods
    you have written so far in this assignment.

    Example test cases:
    b"hello"  -> b"68656c6c6f"
    b"world"  -> b"776f726c64"
    b"the"    -> b"746865"
    b"puzzle" -> b"70757a7a6c65"
    """
    import binascii
    
    # todo: implement this function
    return hexlify(the_input)
```

```python
grader.check("q8")
```

**Question 9 (Hex to Base64).** Write a function `hexstring_to_base64` that receives a hexadecimal encoding of a bytestring, and returns the Base 64 encoding of the same bytestring.

(Note: this is actually the first challenge in the [Cryptopals](https://cryptopals.com/sets/1/challenges/1) crypto coding challenges. We might see some other challenges from Cryptopals later in the semester.)

This time, you might find the `base64` library to be useful. Whenever you explore a new Python library in this course, it is always a good idea to take a few minutes to read its documentation. Here's the documentation for `base64`: https://docs.python.org/3/library/base64.html

```python
import base64
import binascii

def hexstring_to_base64(the_input):
    """ Convert the_input from hex to base64 format.

    Example test cases:
    b'1de965a3ef96a2b95d' -> b'Hello++World'
    b'49276d206b696c6c696e6720796f757220627261696e206c696b65206120706f69736f6e6f7573206d757368726f6f6d' -> b'SSdtIGtpbGxpbmcgeW91ciBicmFpbiBsaWtlIGEgcG9pc29ub3VzIG11c2hyb29t'
    """
    # todo: implement the rest of this function
    return base64.b64encode(unhexlify(the_input))

```

```python
grader.check("q9")
```

**Question 10 (Hashing short strings).** A popular cryptographic hash function is SHA-256. I ran the following code:

```
from hashlib import sha256
bytes_character = # redacted from code
hexlify(sha256(bytes_character).digest())
```

I won't tell you exactly what I chose for the variable `bytes_character` but I will give you a hint: I set this variable to a bytestring of length 1, and its ASCII representation is a single lowercase letter.

The output was: `454349e422f05297191ead13e21d3db520e5abef52055e4964b82fb213f593a1`

Based on this information, what was my input `bytes_character`?

```python
from hashlib import sha256
bytes_character = b'r'
hexlify(sha256(bytes_character).digest())
```

Output:
```
b'454349e422f05297191ead13e21d3db520e5abef52055e4964b82fb213f593a1'
```

```python
def hw1_q10() -> bytes:
    # calculate and return the correct value for bytes_character
    bytes_character = b'r'
    return bytes_character
```

```python
grader.check("q10")
```

### Collaboration Policy

Congratulations on completing the assignment! Before you submit: list all collaborators, sources, and AI tools in accordance with the collaboration policy. Doing so is to your advantage: this is your opportunity to document and explain any external information used and why you believe it adheres to the code of conduct and collaboration policy. If I discover an undocumented violation of the collaboration policy, then this will be considered academic misconduct.

A. Write the names of all classmates you worked with, along with a short description of the work that you performed together.

**Your response:** 

B. List all written materials that you used, such as books or websites (besides the lecture notes and course textbooks). Provide links to any web-based resources, or citations to any physical works.

**Your response:** 

C. State all code that you used from other sources. In particular, if you used an AI tool, then you must include the entire exchange with the AI tool, as per the [CDS Generative AI Assistance Policy](https://www.bu.edu/cds-faculty/culture-community/gaia-policy/).
https://www.geeksforgeeks.org/python-hex-function/

**Your response:** 

### Sending to Gradescope

After completing the assignment, submit only the `.ipynb` file to Gradescope. It takes a while for the auto grading system to check your work.

Tests for q1 (points: 2)
```
>>> from hashlib import sha256
>>> def test_hw1_q1_new():
...     assert sha256(hw1_q1().to_bytes(1, 'big') + b'c0535e4be2b79ffd93291305436bf889').hexdigest() == 'c7f87adb56101dbfdb495c58ea3996e5d6b8bac3969c3f5ae155b38420e5492f'
...     return
>>> test_hw1_q1_new()
```
failure_message: Try again
success_message: Satisfactory


Tests for q10 (points: 3)
```
>>> from hashlib import sha256
>>> def test_hw1_q10():
...     assert sha256(hw1_q10() + b'e66cef3c769831a9351b27971508aae5').hexdigest() == 'e319873bb4c0a5f7b20403d96f0bb71d13bfc4cdfe29c081a4d651f726f441ee'
...     return
>>> test_hw1_q10()
```
failure_message: Try again
success_message: Satisfactory


Tests for q2 (points: 2)
```
>>> from hashlib import sha256
>>> def test_hw1_q2():
...     assert sha256(hw1_q2().to_bytes(1, 'big') + b'045babdcd2118960e8c8b8e0ecf69e7a').hexdigest() == 'c0b1176d546091aaecce38bf59aa71e091e0e4d7536a05ab6503b9f17df60942'
...     return
>>> test_hw1_q2()
```
failure_message: Try again
success_message: Satisfactory


Tests for q3 (points: 2)
```
>>> from hashlib import sha256
>>> def test_hw1_q3():
...     assert sha256(hw1_q3().to_bytes(1, 'big') + b'4ad30843f6b867b8ca49f9f22caaab82').hexdigest() == '1d396f85f99a221d4feb7401f19d519057ff4faa441fc1940c207dd0fb8f5920'
...     return
>>> test_hw1_q3()
```
failure_message: Try again
success_message: Satisfactory


Tests for q4 (points: 2)
```
>>> from hashlib import sha256
>>> def test_hw1_q4():
...     assert sha256(hw1_q4().to_bytes(1, 'big') + b'5d9501fe319a8ec463fdf4325a2a5df4').hexdigest() == '3d5cb89cb7075dee0cef6c26991d6920917e945bcc37f87ca2b858d279262769'
...     return
>>> test_hw1_q4()
```
failure_message: Try again
success_message: Satisfactory


Tests for q5 (points: 2)
```
>>> from hashlib import sha256
>>> def test_hw1_q5():
...     assert sha256(hw1_q5().to_bytes(1, 'big') + b'a339f066967c0000da3fcda4fd7eb4ef').hexdigest() == 'ffdd3df3934dea52a36bdec1e351bfa803030907133cb49d07cab9b4478a402d'
...     return
>>> test_hw1_q5()
```
failure_message: Try again
success_message: Satisfactory


Tests for q6 (points: 2)
```
>>> from hashlib import sha256
>>> def test_hw1_q6():
...     assert sha256(hw1_q6().to_bytes(1, 'big') + b'd0166f6a411c36832e36ae33c07856ce').hexdigest() == '0d1d861aa55f3ea69ae500c9a9eca34762768d2b4a6476ea5026e0f7d0d33055'
...     return
>>> test_hw1_q6()
```
failure_message: Try again
success_message: Satisfactory


Tests for q7 (points: 3)
```
>>> def test_hw1_q7():
...     assert int_byte_to_hex(0) == '00'
...     assert int_byte_to_hex(1) == '01'
...     assert int_byte_to_hex(2) == '02'
...     assert int_byte_to_hex(3) == '03'
...     assert int_byte_to_hex(4) == '04'
...     assert int_byte_to_hex(5) == '05'
...     assert int_byte_to_hex(6) == '06'
...     assert int_byte_to_hex(7) == '07'
...     assert int_byte_to_hex(8) == '08'
...     assert int_byte_to_hex(9) == '09'
...     assert int_byte_to_hex(10) == '0a'
...     assert int_byte_to_hex(11) == '0b'
...     assert int_byte_to_hex(12) == '0c'
...     assert int_byte_to_hex(13) == '0d'
...     assert int_byte_to_hex(14) == '0e'
...     assert int_byte_to_hex(15) == '0f'
...     assert int_byte_to_hex(16) == '10'
...     assert int_byte_to_hex(17) == '11'
...     assert int_byte_to_hex(18) == '12'
...     assert int_byte_to_hex(19) == '13'
...     assert int_byte_to_hex(20) == '14'
...     assert int_byte_to_hex(21) == '15'
...     assert int_byte_to_hex(22) == '16'
...     assert int_byte_to_hex(23) == '17'
...     assert int_byte_to_hex(24) == '18'
...     assert int_byte_to_hex(25) == '19'
...     assert int_byte_to_hex(26) == '1a'
...     assert int_byte_to_hex(27) == '1b'
...     assert int_byte_to_hex(28) == '1c'
...     assert int_byte_to_hex(29) == '1d'
...     assert int_byte_to_hex(30) == '1e'
...     assert int_byte_to_hex(31) == '1f'
...     assert int_byte_to_hex(32) == '20'
...     assert int_byte_to_hex(33) == '21'
...     assert int_byte_to_hex(34) == '22'
...     assert int_byte_to_hex(35) == '23'
...     assert int_byte_to_hex(36) == '24'
...     assert int_byte_to_hex(37) == '25'
...     assert int_byte_to_hex(38) == '26'
...     assert int_byte_to_hex(39) == '27'
...     assert int_byte_to_hex(40) == '28'
...     assert int_byte_to_hex(41) == '29'
...     assert int_byte_to_hex(42) == '2a'
...     assert int_byte_to_hex(43) == '2b'
...     assert int_byte_to_hex(44) == '2c'
...     assert int_byte_to_hex(45) == '2d'
...     assert int_byte_to_hex(46) == '2e'
...     assert int_byte_to_hex(47) == '2f'
...     assert int_byte_to_hex(48) == '30'
...     assert int_byte_to_hex(49) == '31'
...     assert int_byte_to_hex(50) == '32'
...     assert int_byte_to_hex(51) == '33'
...     assert int_byte_to_hex(52) == '34'
...     assert int_byte_to_hex(53) == '35'
...     assert int_byte_to_hex(54) == '36'
...     assert int_byte_to_hex(55) == '37'
...     assert int_byte_to_hex(56) == '38'
...     assert int_byte_to_hex(57) == '39'
...     assert int_byte_to_hex(58) == '3a'
...     assert int_byte_to_hex(59) == '3b'
...     assert int_byte_to_hex(60) == '3c'
...     assert int_byte_to_hex(61) == '3d'
...     assert int_byte_to_hex(62) == '3e'
...     assert int_byte_to_hex(63) == '3f'
...     assert int_byte_to_hex(64) == '40'
...     assert int_byte_to_hex(65) == '41'
...     assert int_byte_to_hex(66) == '42'
...     assert int_byte_to_hex(67) == '43'
...     assert int_byte_to_hex(68) == '44'
...     assert int_byte_to_hex(69) == '45'
...     assert int_byte_to_hex(70) == '46'
...     assert int_byte_to_hex(71) == '47'
...     assert int_byte_to_hex(72) == '48'
...     assert int_byte_to_hex(73) == '49'
...     assert int_byte_to_hex(74) == '4a'
...     assert int_byte_to_hex(75) == '4b'
...     assert int_byte_to_hex(76) == '4c'
...     assert int_byte_to_hex(77) == '4d'
...     assert int_byte_to_hex(78) == '4e'
...     assert int_byte_to_hex(79) == '4f'
...     assert int_byte_to_hex(80) == '50'
...     assert int_byte_to_hex(81) == '51'
...     assert int_byte_to_hex(82) == '52'
...     assert int_byte_to_hex(83) == '53'
...     assert int_byte_to_hex(84) == '54'
...     assert int_byte_to_hex(85) == '55'
...     assert int_byte_to_hex(86) == '56'
...     assert int_byte_to_hex(87) == '57'
...     assert int_byte_to_hex(88) == '58'
...     assert int_byte_to_hex(89) == '59'
...     assert int_byte_to_hex(90) == '5a'
...     assert int_byte_to_hex(91) == '5b'
...     assert int_byte_to_hex(92) == '5c'
...     assert int_byte_to_hex(93) == '5d'
...     assert int_byte_to_hex(94) == '5e'
...     assert int_byte_to_hex(95) == '5f'
...     assert int_byte_to_hex(96) == '60'
...     assert int_byte_to_hex(97) == '61'
...     assert int_byte_to_hex(98) == '62'
...     assert int_byte_to_hex(99) == '63'
...     assert int_byte_to_hex(100) == '64'
...     assert int_byte_to_hex(101) == '65'
...     assert int_byte_to_hex(102) == '66'
...     assert int_byte_to_hex(103) == '67'
...     assert int_byte_to_hex(104) == '68'
...     assert int_byte_to_hex(105) == '69'
...     assert int_byte_to_hex(106) == '6a'
...     assert int_byte_to_hex(107) == '6b'
...     assert int_byte_to_hex(108) == '6c'
...     assert int_byte_to_hex(109) == '6d'
...     assert int_byte_to_hex(110) == '6e'
...     assert int_byte_to_hex(111) == '6f'
...     assert int_byte_to_hex(112) == '70'
...     assert int_byte_to_hex(113) == '71'
...     assert int_byte_to_hex(114) == '72'
...     assert int_byte_to_hex(115) == '73'
...     assert int_byte_to_hex(116) == '74'
...     assert int_byte_to_hex(117) == '75'
...     assert int_byte_to_hex(118) == '76'
...     assert int_byte_to_hex(119) == '77'
...     assert int_byte_to_hex(120) == '78'
...     assert int_byte_to_hex(121) == '79'
...     assert int_byte_to_hex(122) == '7a'
...     assert int_byte_to_hex(123) == '7b'
...     assert int_byte_to_hex(124) == '7c'
...     assert int_byte_to_hex(125) == '7d'
...     assert int_byte_to_hex(126) == '7e'
...     assert int_byte_to_hex(127) == '7f'
...     assert int_byte_to_hex(128) == '80'
...     assert int_byte_to_hex(129) == '81'
...     assert int_byte_to_hex(130) == '82'
...     assert int_byte_to_hex(131) == '83'
...     assert int_byte_to_hex(132) == '84'
...     assert int_byte_to_hex(133) == '85'
...     assert int_byte_to_hex(134) == '86'
...     assert int_byte_to_hex(135) == '87'
...     assert int_byte_to_hex(136) == '88'
...     assert int_byte_to_hex(137) == '89'
...     assert int_byte_to_hex(138) == '8a'
...     assert int_byte_to_hex(139) == '8b'
...     assert int_byte_to_hex(140) == '8c'
...     assert int_byte_to_hex(141) == '8d'
...     assert int_byte_to_hex(142) == '8e'
...     assert int_byte_to_hex(143) == '8f'
...     assert int_byte_to_hex(144) == '90'
...     assert int_byte_to_hex(145) == '91'
...     assert int_byte_to_hex(146) == '92'
...     assert int_byte_to_hex(147) == '93'
...     assert int_byte_to_hex(148) == '94'
...     assert int_byte_to_hex(149) == '95'
...     assert int_byte_to_hex(150) == '96'
...     assert int_byte_to_hex(151) == '97'
...     assert int_byte_to_hex(152) == '98'
...     assert int_byte_to_hex(153) == '99'
...     assert int_byte_to_hex(154) == '9a'
...     assert int_byte_to_hex(155) == '9b'
...     assert int_byte_to_hex(156) == '9c'
...     assert int_byte_to_hex(157) == '9d'
...     assert int_byte_to_hex(158) == '9e'
...     assert int_byte_to_hex(159) == '9f'
...     assert int_byte_to_hex(160) == 'a0'
...     assert int_byte_to_hex(161) == 'a1'
...     assert int_byte_to_hex(162) == 'a2'
...     assert int_byte_to_hex(163) == 'a3'
...     assert int_byte_to_hex(164) == 'a4'
...     assert int_byte_to_hex(165) == 'a5'
...     assert int_byte_to_hex(166) == 'a6'
...     assert int_byte_to_hex(167) == 'a7'
...     assert int_byte_to_hex(168) == 'a8'
...     assert int_byte_to_hex(169) == 'a9'
...     assert int_byte_to_hex(170) == 'aa'
...     assert int_byte_to_hex(171) == 'ab'
...     assert int_byte_to_hex(172) == 'ac'
...     assert int_byte_to_hex(173) == 'ad'
...     assert int_byte_to_hex(174) == 'ae'
...     assert int_byte_to_hex(175) == 'af'
...     assert int_byte_to_hex(176) == 'b0'
...     assert int_byte_to_hex(177) == 'b1'
...     assert int_byte_to_hex(178) == 'b2'
...     assert int_byte_to_hex(179) == 'b3'
...     assert int_byte_to_hex(180) == 'b4'
...     assert int_byte_to_hex(181) == 'b5'
...     assert int_byte_to_hex(182) == 'b6'
...     assert int_byte_to_hex(183) == 'b7'
...     assert int_byte_to_hex(184) == 'b8'
...     assert int_byte_to_hex(185) == 'b9'
...     assert int_byte_to_hex(186) == 'ba'
...     assert int_byte_to_hex(187) == 'bb'
...     assert int_byte_to_hex(188) == 'bc'
...     assert int_byte_to_hex(189) == 'bd'
...     assert int_byte_to_hex(190) == 'be'
...     assert int_byte_to_hex(191) == 'bf'
...     assert int_byte_to_hex(192) == 'c0'
...     assert int_byte_to_hex(193) == 'c1'
...     assert int_byte_to_hex(194) == 'c2'
...     assert int_byte_to_hex(195) == 'c3'
...     assert int_byte_to_hex(196) == 'c4'
...     assert int_byte_to_hex(197) == 'c5'
...     assert int_byte_to_hex(198) == 'c6'
...     assert int_byte_to_hex(199) == 'c7'
...     assert int_byte_to_hex(200) == 'c8'
...     assert int_byte_to_hex(201) == 'c9'
...     assert int_byte_to_hex(202) == 'ca'
...     assert int_byte_to_hex(203) == 'cb'
...     assert int_byte_to_hex(204) == 'cc'
...     assert int_byte_to_hex(205) == 'cd'
...     assert int_byte_to_hex(206) == 'ce'
...     assert int_byte_to_hex(207) == 'cf'
...     assert int_byte_to_hex(208) == 'd0'
...     assert int_byte_to_hex(209) == 'd1'
...     assert int_byte_to_hex(210) == 'd2'
...     assert int_byte_to_hex(211) == 'd3'
...     assert int_byte_to_hex(212) == 'd4'
...     assert int_byte_to_hex(213) == 'd5'
...     assert int_byte_to_hex(214) == 'd6'
...     assert int_byte_to_hex(215) == 'd7'
...     assert int_byte_to_hex(216) == 'd8'
...     assert int_byte_to_hex(217) == 'd9'
...     assert int_byte_to_hex(218) == 'da'
...     assert int_byte_to_hex(219) == 'db'
...     assert int_byte_to_hex(220) == 'dc'
...     assert int_byte_to_hex(221) == 'dd'
...     assert int_byte_to_hex(222) == 'de'
...     assert int_byte_to_hex(223) == 'df'
...     assert int_byte_to_hex(224) == 'e0'
...     assert int_byte_to_hex(225) == 'e1'
...     assert int_byte_to_hex(226) == 'e2'
...     assert int_byte_to_hex(227) == 'e3'
...     assert int_byte_to_hex(228) == 'e4'
...     assert int_byte_to_hex(229) == 'e5'
...     assert int_byte_to_hex(230) == 'e6'
...     assert int_byte_to_hex(231) == 'e7'
...     assert int_byte_to_hex(232) == 'e8'
...     assert int_byte_to_hex(233) == 'e9'
...     assert int_byte_to_hex(234) == 'ea'
...     assert int_byte_to_hex(235) == 'eb'
...     assert int_byte_to_hex(236) == 'ec'
...     assert int_byte_to_hex(237) == 'ed'
...     assert int_byte_to_hex(238) == 'ee'
...     assert int_byte_to_hex(239) == 'ef'
...     assert int_byte_to_hex(240) == 'f0'
...     assert int_byte_to_hex(241) == 'f1'
...     assert int_byte_to_hex(242) == 'f2'
...     assert int_byte_to_hex(243) == 'f3'
...     assert int_byte_to_hex(244) == 'f4'
...     assert int_byte_to_hex(245) == 'f5'
...     assert int_byte_to_hex(246) == 'f6'
...     assert int_byte_to_hex(247) == 'f7'
...     assert int_byte_to_hex(248) == 'f8'
...     assert int_byte_to_hex(249) == 'f9'
...     assert int_byte_to_hex(250) == 'fa'
...     assert int_byte_to_hex(251) == 'fb'
...     assert int_byte_to_hex(252) == 'fc'
...     assert int_byte_to_hex(253) == 'fd'
...     assert int_byte_to_hex(254) == 'fe'
...     assert int_byte_to_hex(255) == 'ff'
...     return
>>> test_hw1_q7()
```
failure_message: Try again
success_message: Satisfactory


Tests for q8 (points: 3)
```
>>> def test_hw1_q8():
...     assert bytestring_to_hex_bytestring(b'ability') == b'6162696c697479'
...     assert bytestring_to_hex_bytestring(b'according') == b'6163636f7264696e67'
...     assert bytestring_to_hex_bytestring(b'act') == b'616374'
...     assert bytestring_to_hex_bytestring(b'activity') == b'6163746976697479'
...     assert bytestring_to_hex_bytestring(b'address') == b'61646472657373'
...     assert bytestring_to_hex_bytestring(b'administration') == b'61646d696e697374726174696f6e'
...     assert bytestring_to_hex_bytestring(b'allow') == b'616c6c6f77'
...     assert bytestring_to_hex_bytestring(b'anything') == b'616e797468696e67'
...     assert bytestring_to_hex_bytestring(b'arrive') == b'617272697665'
...     assert bytestring_to_hex_bytestring(b'bad') == b'626164'
...     assert bytestring_to_hex_bytestring(b'behavior') == b'6265686176696f72'
...     assert bytestring_to_hex_bytestring(b'believe') == b'62656c69657665'
...     assert bytestring_to_hex_bytestring(b'best') == b'62657374'
...     assert bytestring_to_hex_bytestring(b'better') == b'626574746572'
...     assert bytestring_to_hex_bytestring(b'car') == b'636172'
...     assert bytestring_to_hex_bytestring(b'cell') == b'63656c6c'
...     assert bytestring_to_hex_bytestring(b'central') == b'63656e7472616c'
...     assert bytestring_to_hex_bytestring(b'change') == b'6368616e6765'
...     assert bytestring_to_hex_bytestring(b'close') == b'636c6f7365'
...     assert bytestring_to_hex_bytestring(b'collection') == b'636f6c6c656374696f6e'
...     assert bytestring_to_hex_bytestring(b'computer') == b'636f6d7075746572'
...     assert bytestring_to_hex_bytestring(b'create') == b'637265617465'
...     assert bytestring_to_hex_bytestring(b'culture') == b'63756c74757265'
...     assert bytestring_to_hex_bytestring(b'determine') == b'64657465726d696e65'
...     assert bytestring_to_hex_bytestring(b'development') == b'646576656c6f706d656e74'
...     assert bytestring_to_hex_bytestring(b'discover') == b'646973636f766572'
...     assert bytestring_to_hex_bytestring(b'down') == b'646f776e'
...     assert bytestring_to_hex_bytestring(b'environment') == b'656e7669726f6e6d656e74'
...     assert bytestring_to_hex_bytestring(b'environmental') == b'656e7669726f6e6d656e74616c'
...     assert bytestring_to_hex_bytestring(b'especially') == b'657370656369616c6c79'
...     assert bytestring_to_hex_bytestring(b'everybody') == b'6576657279626f6479'
...     assert bytestring_to_hex_bytestring(b'everyone') == b'65766572796f6e65'
...     assert bytestring_to_hex_bytestring(b'exist') == b'6578697374'
...     assert bytestring_to_hex_bytestring(b'expect') == b'657870656374'
...     assert bytestring_to_hex_bytestring(b'federal') == b'6665646572616c'
...     assert bytestring_to_hex_bytestring(b'fly') == b'666c79'
...     assert bytestring_to_hex_bytestring(b'follow') == b'666f6c6c6f77'
...     assert bytestring_to_hex_bytestring(b'front') == b'66726f6e74'
...     assert bytestring_to_hex_bytestring(b'happen') == b'68617070656e'
...     assert bytestring_to_hex_bytestring(b'heavy') == b'6865617679'
...     assert bytestring_to_hex_bytestring(b'hospital') == b'686f73706974616c'
...     assert bytestring_to_hex_bytestring(b'hour') == b'686f7572'
...     assert bytestring_to_hex_bytestring(b'how') == b'686f77'
...     assert bytestring_to_hex_bytestring(b'imagine') == b'696d6167696e65'
...     assert bytestring_to_hex_bytestring(b'increase') == b'696e637265617365'
...     assert bytestring_to_hex_bytestring(b'information') == b'696e666f726d6174696f6e'
...     assert bytestring_to_hex_bytestring(b'institution') == b'696e737469747574696f6e'
...     assert bytestring_to_hex_bytestring(b'interesting') == b'696e746572657374696e67'
...     assert bytestring_to_hex_bytestring(b'international') == b'696e7465726e6174696f6e616c'
...     assert bytestring_to_hex_bytestring(b'investment') == b'696e766573746d656e74'
...     assert bytestring_to_hex_bytestring(b'knowledge') == b'6b6e6f776c65646765'
...     assert bytestring_to_hex_bytestring(b'land') == b'6c616e64'
...     assert bytestring_to_hex_bytestring(b'language') == b'6c616e6775616765'
...     assert bytestring_to_hex_bytestring(b'laugh') == b'6c61756768'
...     assert bytestring_to_hex_bytestring(b'letter') == b'6c6574746572'
...     assert bytestring_to_hex_bytestring(b'low') == b'6c6f77'
...     assert bytestring_to_hex_bytestring(b'machine') == b'6d616368696e65'
...     assert bytestring_to_hex_bytestring(b'maintain') == b'6d61696e7461696e'
...     assert bytestring_to_hex_bytestring(b'memory') == b'6d656d6f7279'
...     assert bytestring_to_hex_bytestring(b'message') == b'6d657373616765'
...     assert bytestring_to_hex_bytestring(b'movie') == b'6d6f766965'
...     assert bytestring_to_hex_bytestring(b'nation') == b'6e6174696f6e'
...     assert bytestring_to_hex_bytestring(b'nothing') == b'6e6f7468696e67'
...     assert bytestring_to_hex_bytestring(b'one') == b'6f6e65'
...     assert bytestring_to_hex_bytestring(b'opportunity') == b'6f70706f7274756e697479'
...     assert bytestring_to_hex_bytestring(b'order') == b'6f72646572'
...     assert bytestring_to_hex_bytestring(b'organization') == b'6f7267616e697a6174696f6e'
...     assert bytestring_to_hex_bytestring(b'painting') == b'7061696e74696e67'
...     assert bytestring_to_hex_bytestring(b'parent') == b'706172656e74'
...     assert bytestring_to_hex_bytestring(b'participant') == b'7061727469636970616e74'
...     assert bytestring_to_hex_bytestring(b'particularly') == b'706172746963756c61726c79'
...     assert bytestring_to_hex_bytestring(b'pass') == b'70617373'
...     assert bytestring_to_hex_bytestring(b'perform') == b'706572666f726d'
...     assert bytestring_to_hex_bytestring(b'performance') == b'706572666f726d616e6365'
...     assert bytestring_to_hex_bytestring(b'place') == b'706c616365'
...     assert bytestring_to_hex_bytestring(b'position') == b'706f736974696f6e'
...     assert bytestring_to_hex_bytestring(b'pressure') == b'7072657373757265'
...     assert bytestring_to_hex_bytestring(b'professional') == b'70726f66657373696f6e616c'
...     assert bytestring_to_hex_bytestring(b'professor') == b'70726f666573736f72'
...     assert bytestring_to_hex_bytestring(b'program') == b'70726f6772616d'
...     assert bytestring_to_hex_bytestring(b'quickly') == b'717569636b6c79'
...     assert bytestring_to_hex_bytestring(b'real') == b'7265616c'
...     assert bytestring_to_hex_bytestring(b'relate') == b'72656c617465'
...     assert bytestring_to_hex_bytestring(b'relationship') == b'72656c6174696f6e73686970'
...     assert bytestring_to_hex_bytestring(b'research') == b'7265736561726368'
...     assert bytestring_to_hex_bytestring(b'respond') == b'726573706f6e64'
...     assert bytestring_to_hex_bytestring(b'responsibility') == b'726573706f6e736962696c697479'
...     assert bytestring_to_hex_bytestring(b'same') == b'73616d65'
...     assert bytestring_to_hex_bytestring(b'season') == b'736561736f6e'
...     assert bytestring_to_hex_bytestring(b'see') == b'736565'
...     assert bytestring_to_hex_bytestring(b'shoulder') == b'73686f756c646572'
...     assert bytestring_to_hex_bytestring(b'significant') == b'7369676e69666963616e74'
...     assert bytestring_to_hex_bytestring(b'sing') == b'73696e67'
...     assert bytestring_to_hex_bytestring(b'small') == b'736d616c6c'
...     assert bytestring_to_hex_bytestring(b'something') == b'736f6d657468696e67'
...     assert bytestring_to_hex_bytestring(b'specific') == b'7370656369666963'
...     assert bytestring_to_hex_bytestring(b'staff') == b'7374616666'
...     assert bytestring_to_hex_bytestring(b'stop') == b'73746f70'
...     assert bytestring_to_hex_bytestring(b'story') == b'73746f7279'
...     assert bytestring_to_hex_bytestring(b'student') == b'73747564656e74'
...     assert bytestring_to_hex_bytestring(b'teacher') == b'74656163686572'
...     assert bytestring_to_hex_bytestring(b'technology') == b'746563686e6f6c6f6779'
...     assert bytestring_to_hex_bytestring(b'than') == b'7468616e'
...     assert bytestring_to_hex_bytestring(b'themselves') == b'7468656d73656c766573'
...     assert bytestring_to_hex_bytestring(b'there') == b'7468657265'
...     assert bytestring_to_hex_bytestring(b'thousand') == b'74686f7573616e64'
...     assert bytestring_to_hex_bytestring(b'today') == b'746f646179'
...     assert bytestring_to_hex_bytestring(b'top') == b'746f70'
...     assert bytestring_to_hex_bytestring(b'traditional') == b'747261646974696f6e616c'
...     assert bytestring_to_hex_bytestring(b'true') == b'74727565'
...     assert bytestring_to_hex_bytestring(b'understand') == b'756e6465727374616e64'
...     assert bytestring_to_hex_bytestring(b'until') == b'756e74696c'
...     assert bytestring_to_hex_bytestring(b'various') == b'766172696f7573'
...     assert bytestring_to_hex_bytestring(b'want') == b'77616e74'
...     assert bytestring_to_hex_bytestring(b'wind') == b'77696e64'
...     assert bytestring_to_hex_bytestring(b'write') == b'7772697465'
...     assert bytestring_to_hex_bytestring(b'you') == b'796f75'
...     return
>>> test_hw1_q8()
```
failure_message: Try again
success_message: Satisfactory


Tests for q9 (points: 3)
```
>>> def test_hw1_q9():
...     assert hexstring_to_base64(b'6162696c697479') == b'YWJpbGl0eQ=='
...     assert hexstring_to_base64(b'6163636f7264696e67') == b'YWNjb3JkaW5n'
...     assert hexstring_to_base64(b'616374') == b'YWN0'
...     assert hexstring_to_base64(b'6163746976697479') == b'YWN0aXZpdHk='
...     assert hexstring_to_base64(b'61646472657373') == b'YWRkcmVzcw=='
...     assert hexstring_to_base64(b'61646d696e697374726174696f6e') == b'YWRtaW5pc3RyYXRpb24='
...     assert hexstring_to_base64(b'616c6c6f77') == b'YWxsb3c='
...     assert hexstring_to_base64(b'616e797468696e67') == b'YW55dGhpbmc='
...     assert hexstring_to_base64(b'617272697665') == b'YXJyaXZl'
...     assert hexstring_to_base64(b'626164') == b'YmFk'
...     assert hexstring_to_base64(b'6265686176696f72') == b'YmVoYXZpb3I='
...     assert hexstring_to_base64(b'62656c69657665') == b'YmVsaWV2ZQ=='
...     assert hexstring_to_base64(b'62657374') == b'YmVzdA=='
...     assert hexstring_to_base64(b'626574746572') == b'YmV0dGVy'
...     assert hexstring_to_base64(b'636172') == b'Y2Fy'
...     assert hexstring_to_base64(b'63656c6c') == b'Y2VsbA=='
...     assert hexstring_to_base64(b'63656e7472616c') == b'Y2VudHJhbA=='
...     assert hexstring_to_base64(b'6368616e6765') == b'Y2hhbmdl'
...     assert hexstring_to_base64(b'636c6f7365') == b'Y2xvc2U='
...     assert hexstring_to_base64(b'636f6c6c656374696f6e') == b'Y29sbGVjdGlvbg=='
...     assert hexstring_to_base64(b'636f6d7075746572') == b'Y29tcHV0ZXI='
...     assert hexstring_to_base64(b'637265617465') == b'Y3JlYXRl'
...     assert hexstring_to_base64(b'63756c74757265') == b'Y3VsdHVyZQ=='
...     assert hexstring_to_base64(b'64657465726d696e65') == b'ZGV0ZXJtaW5l'
...     assert hexstring_to_base64(b'646576656c6f706d656e74') == b'ZGV2ZWxvcG1lbnQ='
...     assert hexstring_to_base64(b'646973636f766572') == b'ZGlzY292ZXI='
...     assert hexstring_to_base64(b'646f776e') == b'ZG93bg=='
...     assert hexstring_to_base64(b'656e7669726f6e6d656e74') == b'ZW52aXJvbm1lbnQ='
...     assert hexstring_to_base64(b'656e7669726f6e6d656e74616c') == b'ZW52aXJvbm1lbnRhbA=='
...     assert hexstring_to_base64(b'657370656369616c6c79') == b'ZXNwZWNpYWxseQ=='
...     assert hexstring_to_base64(b'6576657279626f6479') == b'ZXZlcnlib2R5'
...     assert hexstring_to_base64(b'65766572796f6e65') == b'ZXZlcnlvbmU='
...     assert hexstring_to_base64(b'6578697374') == b'ZXhpc3Q='
...     assert hexstring_to_base64(b'657870656374') == b'ZXhwZWN0'
...     assert hexstring_to_base64(b'6665646572616c') == b'ZmVkZXJhbA=='
...     assert hexstring_to_base64(b'666c79') == b'Zmx5'
...     assert hexstring_to_base64(b'666f6c6c6f77') == b'Zm9sbG93'
...     assert hexstring_to_base64(b'66726f6e74') == b'ZnJvbnQ='
...     assert hexstring_to_base64(b'68617070656e') == b'aGFwcGVu'
...     assert hexstring_to_base64(b'6865617679') == b'aGVhdnk='
...     assert hexstring_to_base64(b'686f73706974616c') == b'aG9zcGl0YWw='
...     assert hexstring_to_base64(b'686f7572') == b'aG91cg=='
...     assert hexstring_to_base64(b'686f77') == b'aG93'
...     assert hexstring_to_base64(b'696d6167696e65') == b'aW1hZ2luZQ=='
...     assert hexstring_to_base64(b'696e637265617365') == b'aW5jcmVhc2U='
...     assert hexstring_to_base64(b'696e666f726d6174696f6e') == b'aW5mb3JtYXRpb24='
...     assert hexstring_to_base64(b'696e737469747574696f6e') == b'aW5zdGl0dXRpb24='
...     assert hexstring_to_base64(b'696e746572657374696e67') == b'aW50ZXJlc3Rpbmc='
...     assert hexstring_to_base64(b'696e7465726e6174696f6e616c') == b'aW50ZXJuYXRpb25hbA=='
...     assert hexstring_to_base64(b'696e766573746d656e74') == b'aW52ZXN0bWVudA=='
...     assert hexstring_to_base64(b'6b6e6f776c65646765') == b'a25vd2xlZGdl'
...     assert hexstring_to_base64(b'6c616e64') == b'bGFuZA=='
...     assert hexstring_to_base64(b'6c616e6775616765') == b'bGFuZ3VhZ2U='
...     assert hexstring_to_base64(b'6c61756768') == b'bGF1Z2g='
...     assert hexstring_to_base64(b'6c6574746572') == b'bGV0dGVy'
...     assert hexstring_to_base64(b'6c6f77') == b'bG93'
...     assert hexstring_to_base64(b'6d616368696e65') == b'bWFjaGluZQ=='
...     assert hexstring_to_base64(b'6d61696e7461696e') == b'bWFpbnRhaW4='
...     assert hexstring_to_base64(b'6d656d6f7279') == b'bWVtb3J5'
...     assert hexstring_to_base64(b'6d657373616765') == b'bWVzc2FnZQ=='
...     assert hexstring_to_base64(b'6d6f766965') == b'bW92aWU='
...     assert hexstring_to_base64(b'6e6174696f6e') == b'bmF0aW9u'
...     assert hexstring_to_base64(b'6e6f7468696e67') == b'bm90aGluZw=='
...     assert hexstring_to_base64(b'6f6e65') == b'b25l'
...     assert hexstring_to_base64(b'6f70706f7274756e697479') == b'b3Bwb3J0dW5pdHk='
...     assert hexstring_to_base64(b'6f72646572') == b'b3JkZXI='
...     assert hexstring_to_base64(b'6f7267616e697a6174696f6e') == b'b3JnYW5pemF0aW9u'
...     assert hexstring_to_base64(b'7061696e74696e67') == b'cGFpbnRpbmc='
...     assert hexstring_to_base64(b'706172656e74') == b'cGFyZW50'
...     assert hexstring_to_base64(b'7061727469636970616e74') == b'cGFydGljaXBhbnQ='
...     assert hexstring_to_base64(b'706172746963756c61726c79') == b'cGFydGljdWxhcmx5'
...     assert hexstring_to_base64(b'70617373') == b'cGFzcw=='
...     assert hexstring_to_base64(b'706572666f726d') == b'cGVyZm9ybQ=='
...     assert hexstring_to_base64(b'706572666f726d616e6365') == b'cGVyZm9ybWFuY2U='
...     assert hexstring_to_base64(b'706c616365') == b'cGxhY2U='
...     assert hexstring_to_base64(b'706f736974696f6e') == b'cG9zaXRpb24='
...     assert hexstring_to_base64(b'7072657373757265') == b'cHJlc3N1cmU='
...     assert hexstring_to_base64(b'70726f66657373696f6e616c') == b'cHJvZmVzc2lvbmFs'
...     assert hexstring_to_base64(b'70726f666573736f72') == b'cHJvZmVzc29y'
...     assert hexstring_to_base64(b'70726f6772616d') == b'cHJvZ3JhbQ=='
...     assert hexstring_to_base64(b'717569636b6c79') == b'cXVpY2tseQ=='
...     assert hexstring_to_base64(b'7265616c') == b'cmVhbA=='
...     assert hexstring_to_base64(b'72656c617465') == b'cmVsYXRl'
...     assert hexstring_to_base64(b'72656c6174696f6e73686970') == b'cmVsYXRpb25zaGlw'
...     assert hexstring_to_base64(b'7265736561726368') == b'cmVzZWFyY2g='
...     assert hexstring_to_base64(b'726573706f6e64') == b'cmVzcG9uZA=='
...     assert hexstring_to_base64(b'726573706f6e736962696c697479') == b'cmVzcG9uc2liaWxpdHk='
...     assert hexstring_to_base64(b'73616d65') == b'c2FtZQ=='
...     assert hexstring_to_base64(b'736561736f6e') == b'c2Vhc29u'
...     assert hexstring_to_base64(b'736565') == b'c2Vl'
...     assert hexstring_to_base64(b'73686f756c646572') == b'c2hvdWxkZXI='
...     assert hexstring_to_base64(b'7369676e69666963616e74') == b'c2lnbmlmaWNhbnQ='
...     assert hexstring_to_base64(b'73696e67') == b'c2luZw=='
...     assert hexstring_to_base64(b'736d616c6c') == b'c21hbGw='
...     assert hexstring_to_base64(b'736f6d657468696e67') == b'c29tZXRoaW5n'
...     assert hexstring_to_base64(b'7370656369666963') == b'c3BlY2lmaWM='
...     assert hexstring_to_base64(b'7374616666') == b'c3RhZmY='
...     assert hexstring_to_base64(b'73746f70') == b'c3RvcA=='
...     assert hexstring_to_base64(b'73746f7279') == b'c3Rvcnk='
...     assert hexstring_to_base64(b'73747564656e74') == b'c3R1ZGVudA=='
...     assert hexstring_to_base64(b'74656163686572') == b'dGVhY2hlcg=='
...     assert hexstring_to_base64(b'746563686e6f6c6f6779') == b'dGVjaG5vbG9neQ=='
...     assert hexstring_to_base64(b'7468616e') == b'dGhhbg=='
...     assert hexstring_to_base64(b'7468656d73656c766573') == b'dGhlbXNlbHZlcw=='
...     assert hexstring_to_base64(b'7468657265') == b'dGhlcmU='
...     assert hexstring_to_base64(b'74686f7573616e64') == b'dGhvdXNhbmQ='
...     assert hexstring_to_base64(b'746f646179') == b'dG9kYXk='
...     assert hexstring_to_base64(b'746f70') == b'dG9w'
...     assert hexstring_to_base64(b'747261646974696f6e616c') == b'dHJhZGl0aW9uYWw='
...     assert hexstring_to_base64(b'74727565') == b'dHJ1ZQ=='
...     assert hexstring_to_base64(b'756e6465727374616e64') == b'dW5kZXJzdGFuZA=='
...     assert hexstring_to_base64(b'756e74696c') == b'dW50aWw='
...     assert hexstring_to_base64(b'766172696f7573') == b'dmFyaW91cw=='
...     assert hexstring_to_base64(b'77616e74') == b'd2FudA=='
...     assert hexstring_to_base64(b'77696e64') == b'd2luZA=='
...     assert hexstring_to_base64(b'7772697465') == b'd3JpdGU='
...     assert hexstring_to_base64(b'796f75') == b'eW91'
...     return
>>> test_hw1_q9()
```
failure_message: Try again
success_message: Satisfactory
